ZC3H11A mutations cause high myopia by triggering PI3K-AKT and NF-κB mediated signaling pathway in humans and mice

Author response:

The following is the authors’ response to the original reviews.

Reviewer #1 (Public Review):

Chen and colleagues investigated ZC3H11A as a potential cause of high myopia (HM) in humans through the analysis of exome sequencing in 1,015 adolescents and experiments involving Zc3h11a knock-out mice. The authors showed four possibly pathogenic missense variants in four adolescents with HM. After that, the authors presented the phenotypic features of Zc3h11a knock-out mice, the result of RNA-sequencing, and a comparison of mRNA and protein levels of the functional candidates between wild-type and Zc3h11a knock-out mice. Based on their observations, the authors concluded that ZC3H11A protein contributes to the early onset of myopia.

The strengths of this manuscript include: (1) successful identification of characteristic ophthalmic phenotypes in Zc3h11a knock-out mice, (2) demonstration of biological features related to myopia, such as PI3K-AKT and NF-kB pathways, and (3) inclusion of supporting human genetic data in individuals with HM. On the other hand, the weaknesses of this paper appear to be: (1) the lack of robust evidence from their genomic analysis, and (2) insufficient evidence to support phenotypic similarity between humans with ZC3H11A mutations and Zc3h11a knock-out mice. Given that the biological mechanisms of high myopia are not fully understood, the identification of a novel gene is valuable. As described in the manuscript, it is worth noting that the previous study using myopic mouse model has implicated the role of ZC3H11A in the etiology of myopia (Fan et al. Plos Genet 2012).

Thank you very much for your valuable suggestions.

Specific comments:

(1) I am concerned about the certainty of similarity in phenotypes between individuals with ZC3H11A mutation and Zc3h11a knock-out mice. A crucial point would be that there are no statistical differences in axial lengths (ALs) between wild-type and Zc3h11a knock-out mice at 8W and 10W, even though ALs in the individuals with ZC3H11A mutation were long. I would also like to note that the phenotypic information of these individuals is not available in the manuscript, although the authors indicated the suppressed b-wave amplitude in Zc3h11a knock-out mice. Considering that the authors described that "Detailed ophthalmic examinations were performed (lines: 321-323)", the detailed clinical features of these individuals should be included in the manuscript.

Thank you for your valuable comments. The axial length in Zc3h11a Het-KO mice were found to be significantly greater than in WT littermates at weeks 4 and 6 (Independent samples t-test, p<0.05; Figure 2A and B). Although no significant differences were observed at other time points, there was still some degree of increase in these parameters. We continued to measure corneal curvature and found no significant differences between the two groups. Therefore, the difference in refraction may be due to the small size of the mouse eye. A 1 D change in refraction corresponds to only a 5-6 μm change in AL(1). However, the SD-OCT resolution used in this study is relatively low (theoretical resolution of 6 μm)(2, 3), so the small changes measured in vitreous cavity depth and AL may not be statistically significant. Additionally, some studies have shown that axial lengths reported in frozen sections are longer than those measured in vivo for age-matched mice(1, 4). Another possible explanation is that the curvature and refractive power of the lens have changed. These hypotheses provide a reasonable explanation for the mismatch between changes in refraction and ocular length parameters.

Reference

(1) Schmucker C, Schaeffel F. A paraxial schematic eye model for the growing C57BL/6 mouse. Vision research 44, 1857-1867 (2004).

(2) Yuan Y, Chen F, Shen M, Lu F, Wang J. Repeated measurements of the anterior segment during accommodation using long scan depth optical coherence tomography. Eye & contact lens 38, 102-108 (2012).

(3) Shen M, et al. SD-OCT with prolonged scan depth for imaging the anterior segment of the eye. Ophthalmic Surgery, Lasers and Imaging Retina 41, S65-S69 (2010).

(4) Schmucker C, Schaeffel F. In vivo biometry in the mouse eye with low coherence interferometry. Vision research 44, 2445-2456 (2004).

Additionally, regarding the “detailed ophthalmic examinations”, due to our patients were selected from a myopia screening cohort of over one million (children and adolescents myopia survey [CAMS] program), and ophthalmic examination only includes semi-annual refractive error measurements (a total of 5 times, with refractive error being the average of the three maximum values) and only one axial length measurement. The inappropriate description of “Detailed clinical features” has been removed.

(2) The term "pathogenic variant" should be used cautiously. Please clarify the pathogenicity of the reported variants in accordance with the ACMG guideline.

Thank you for your valuable comments. Four missense mutations in the ZC3H11A gene (c.412G>A, p.V138I; c.128G>A, p.G43E; c.461C>T, p.P154L; and c.2239T>A, p.S747T) were identified in the 1015 HM patients aged from 15 to 18 years. All of the identified mutations exhibited very low frequencies or does not exist in the Genome Aggregation Database (gnomAD) and Clinvar, and using pathogenicity prediction software SIFT, PolyPhen2, and CADD, most of them display high pathogenicity levels. Among them, c.412G>A, c.128G>A and c.461C>T were located in or around a domain named zf-CCCH_3 (Figure 1A and B). Furthermore, all of the mutation sites were located in highly conserved amino acids across different species (Figure 1C). Four mutations resulted in a higher degree of conformational flexibility and altered the negative charge at the corresponding sites (Figure 1D and E). Meanwhile, through transfection of overexpression mutant plasmids, it was found that compared to the wild-type, the mRNA expression levels of IκBα in the nucleus of all four mutant types (ZC3H11AV138I, ZC3H11AG43E, ZC3H11AP154L and ZC3H11AS747T) were significantly reduced (Supplement Figure 3). According to the ACMG guidelines, the above mutations can be classified as “Pathogenic Moderate”.

(3) The genetic analysis does not fully support the claim that ZC3H11A is causative for HM. While the authors showed the rare allele frequencies and high CADD scores (> 20) of the identified variants, these were insufficient to establish causality. A helpful way to assess the causality would be performing a segregation analysis. An alternative approach is to show significant association by performing a gene-level association test. Assessing the pathogenicity of the variants using various prediction software, such as SIFT, PolyPhen2, and REVEL may also provide additional supportive evidence.

Thank you for your valuable comments. We have addad the pathogenicity of the variants using various prediction software, such as SIFT, PolyPhen2, CADD, and the population variation databases, such as Genome Aggregation Database (gnomAD_AF) and ClinVar. Meanwhile, through transfection of overexpression mutant plasmids, it was found that compared to the wild-type, the mRNA expression levels of IκBα in the nucleus of all four mutant types (ZC3H11AV138I, ZC3H11AG43E, ZC3H11AP154L and ZC3H11AS747T) were significantly reduced (Supplement Figure 3).

(4) As shown in Figure 2, significant differences in refraction were observed from 4 weeks to 10 weeks. Nevertheless, no differences were observed in AL, anterior/vitreous chamber depth, and lens depth. The author should experimentally clarify what factors contribute to the observed difference in refraction.

Thank you for your valuable comments. The existing data show significant differences in refraction between 4 and 10 weeks, with the AL and vitreous cavity depth of Het mice being longer than those of WT mice at 4 and 6 weeks. Although no significant differences were observed at other time points, there was still some degree of increase in these parameters. We continued to measure corneal curvature and found no significant differences between the two groups. Therefore, the difference in refraction may be due to the small size of the mouse eye. A 1 D change in refraction corresponds to only a 5-6 μm change in AL(1). However, the SD-OCT resolution used in this study is relatively low (theoretical resolution of 6 μm)(2, 3), so the small changes measured in vitreous cavity depth and AL may not be statistically significant. Additionally, some studies have shown that axial lengths reported in frozen sections are longer than those measured in vivo for age-matched mice(1, 4). Another possible explanation is that the curvature and refractive power of the lens have changed. These hypotheses provide a reasonable explanation for the mismatch between changes in refraction and ocular length parameters.

Reference

(1) Schmucker C, Schaeffel F. A paraxial schematic eye model for the growing C57BL/6 mouse. Vision research 44, 1857-1867 (2004).

(2) Yuan Y, Chen F, Shen M, Lu F, Wang J. Repeated measurements of the anterior segment during accommodation using long scan depth optical coherence tomography. Eye & contact lens 38, 102-108 (2012).

(3) Shen M, et al. SD-OCT with prolonged scan depth for imaging the anterior segment of the eye. Ophthalmic Surgery, Lasers and Imaging Retina 41, S65-S69 (2010).

(4) Schmucker C, Schaeffel F. In vivo biometry in the mouse eye with low coherence interferometry. Vision research 44, 2445-2456 (2004).

(5) The gene names should be italicized throughout the manuscript.

Thank you for your valuable comments. The gene names have been italicized throughout the manuscript.

(6) Table 1: providing chromosomal positions and rs numbers (if available) would be helpful for readers.

Thank you for your valuable comments. We have provided the chromosome positions and rs number (if available) of each mutation in Table 1.

(7) Figure 5b, c, and d: the results of pathway analysis and GO enrichment analysis are difficult to interpret due to the small font size. It would be preferable to present these results in tables. Moreover, the authors should set a significant threshold in the enrichment analyses.

Thank you for your valuable comments. We have adjusted the font size of the image. In the retina transcriptome analysis, we have set Fold change (FC) of at least two and a P value < 0.05 as thresholds to analyze differentially expressed genes (DEGs). The GO terms and KEGG pathways enrichment analysis selected the top 20 with the most significant differences or the highest number of enriched genes for display.

Reviewer #2 (Public Review):

Summary: Chong Chen and colleagues reported that mutations were identified in the ZC3H11A gene in four adolescents from 1015 high myopia subjects in their myopia cohort. They further generated Zc3h11a knockout mice utilizing the CRISPR/Cas9 technology. They analyzed the heterozygotes knockout mice compared to control littermates and found refractive error changes, electrophysiological differences, and retinal inflammation-related gene expression differences. They concluded that ZC3H11A may play a role in the early onset of myopia by regulating inflammatory responses.

Strengths:

Data were shown from both clinical cohort and animal models.

Weaknesses:

Their findings are interesting and important, however; they need to resolve several points to make the current conclusion.

(1) They described the ZC3H11A gene as a pathogenic variant for high myopia. It should be classified as pathogenic according to the guidelines of the American College of Medical Genetics and Genomics (Richards et al., Genet Med 17(5):405-24, 2015). The modes of inheritance for the families need to be shown. They also described identifying the gene as a "new" candidate. It should be checked in databases such as gnomAD and ClinVar, and any previous publications and be declared as a novel variant.

Thank you for your valuable comments. Four missense mutations in the ZC3H11A gene (c.412G>A, p.V138I; c.128G>A, p.G43E; c.461C>T, p.P154L; and c.2239T>A, p.S747T) were identified in the 1015 HM patients aged from 15 to 18 years. All of the identified mutations exhibited very low frequencies or does not exist in the Genome Aggregation Database (gnomAD) and Clinvar, and using pathogenicity prediction software SIFT, PolyPhen2, and CADD, most of them display high pathogenicity levels. Among them, c.412G>A, c.128G>A and c.461C>T were located in or around a domain named zf-CCCH_3 (Figure 1A and B). Furthermore, all of the mutation sites were located in highly conserved amino acids across different species (Figure 1C). Four mutations resulted in a higher degree of conformational flexibility and altered the negative charge at the corresponding sites (Figure 1D and E). Meanwhile, through transfection of overexpression mutant plasmids, it was found that compared to the wild-type, the mRNA expression levels of IκBα in the nucleus of all four mutant types (ZC3H11AV138I, ZC3H11AG43E, ZC3H11AP154L and ZC3H11AS747T) were significantly reduced (Supplement Figure 3). According to the ACMG guidelines, the above mutations can be classified as “Pathogenic Moderate”.

Unfortunately, our patients are part of the MAGIC project (aged 15 years or older), a cohort consists of thousands of individuals with HM (patients from the children and adolescents myopia survey [CAMS] program) who have undergone WES, and their parents' relevant information was not collected for performing a segregation analysis.

(2) The phenotypes of the heterozygote mice are weak overall. The het mice showed mild to moderate myopic refractive shifts from 4 to 10 weeks of age. However, this cannot be explained by other ocular biometrics such as anterior chamber depth or lens thickness. Some differences are found between het and WT littermates in axial length and vitreous chamber depth but disappear after 8 weeks old. Furthermore, the early differences are not enough to explain the refractive error changes. They mentioned that they did not use homozygotes because of the embryonic lethality. I would strongly suggest employing conditional knockout systems to analyze homozygotes. This will also be able to identify the causative tissues/cells because they assume bipolar cells are functional. The cells in the retinal pigment epithelium and choroid are also important to contribute to myopia development.

Thank you for your valuable comments. The existing data show significant differences in refraction between 4 and 10 weeks, with the AL and vitreous cavity depth of Het mice being longer than those of WT mice at 4 and 6 weeks. Although no significant differences were observed at other time points, there was still some degree of increase in these parameters. We continued to measure corneal curvature and found no significant differences between the two groups. Therefore, the difference in refraction may be due to the small size of the mouse eye. A 1 D change in refraction corresponds to only a 5-6 μm change in AL(1). However, the SD-OCT resolution used in this study is relatively low (theoretical resolution of 6 μm)(2, 3), so the small changes measured in vitreous cavity depth and AL may not be statistically significant. Additionally, some studies have shown that axial lengths reported in frozen sections are longer than those measured in vivo for age-matched mice(1, 4). Another possible explanation is that the curvature and refractive power of the lens have changed. These hypotheses provide a reasonable explanation for the mismatch between changes in refraction and ocular length parameters.

Reference

(1) Schmucker C, Schaeffel F. A paraxial schematic eye model for the growing C57BL/6 mouse. Vision research 44, 1857-1867 (2004).

(2) Yuan Y, Chen F, Shen M, Lu F, Wang J. Repeated measurements of the anterior segment during accommodation using long scan depth optical coherence tomography. Eye & contact lens 38, 102-108 (2012).

(3) Shen M, et al. SD-OCT with prolonged scan depth for imaging the anterior segment of the eye. Ophthalmic Surgery, Lasers and Imaging Retina 41, S65-S69 (2010).

(4) Schmucker C, Schaeffel F. In vivo biometry in the mouse eye with low coherence interferometry. Vision research 44, 2445-2456 (2004).

The drawback is that, we did not conduct relevant research on homozygous knockout mice. The first reason is that our patient's mutation pattern is heterozygous mutation (Heterozygous knockout mice can better simulate human phenotypes). The second reason is that homozygous knockout mice are lethal, and we did not use the conditional knockout mouse model for further research. At the same time, we limited the pathway of myopia to the recognized and classical retina-sclera pathway, and did not study other pathways such as retinal pigment epithelium and choroid.

(3) Their hypothesis regarding inflammatory gene changes and myopic development is not logical. Are the inflammatory responses evoked from bipolar cells? Did the mice show an accumulation of inflammatory cells in the inner retina? Visible retinal inflammation is not generally seen in either early-onset or high-myopia human subjects. Can this be seen in the actual subjects in the cohort? To me, this is difficult to adapt the retina-to-sclera signaling they mentioned in the discussion so far. Egr-1 may be examined as described.

Thank you for your valuable comments. We have removed the hypothesis regarding inflammatory gene changes and myopic development. At present, the explanation is based solely on the correlation of signal pathways, the theoretical basis comes from the reference literature:

“Lin et al., Role of Chronic Inflammation in Myopia Progression: Clinical Evidence and Experimental Validation. EBioMedicine, 2016 Aug:10:269-81, Figure 7.”

Reviewer #3 (Public Review):

Chen et al have identified a new candidate gene for high myopia, ZC3H11A, and using a knock-out mouse model, have attempted to validate it as a myopia gene and explain a potential mechanism. They identified 4 heterozygous missense variants in highly myopic teenagers. These variants are in conserved regions of the protein, but the authors provide no evidence that these specific variants affect protein function. They then created a knock-out mouse. Heterozygotes show myopia at all ages examined but increased axial length only at very early ages. Unfortunately, the authors do not address this point or examine corneal structure in these animals. They show that the mice have decreased B-wave amplitude on electroretinogram (a sign of retinal dysfunction associated with bipolar cells), and decreased expression of a bipolar cell marker, PKCa. They do not address, however, whether there are fewer bipolar cells, or simply decreased expression of the marker protein. On electron microscopy, there are morphologic differences in the outer nuclear layer (where bipolar, amacrine, and horizontal cell bodies reside). Transcriptome analysis identified over 700 differentially expressed genes. The authors chose to focus on the PI3K-AKT and NF-kB signaling pathways and show changes in the expression of genes and proteins in those pathways, including PI3K, AKT, IkBa, NF-kB, TGF-b1, MMP-2, and IL-6, although there is very high variability between animals. They propose that myopia may develop in these animals either as a result of visual abnormality (decreased bipolar cell function in the retina) or by alteration of NF-kB signaling. These data provide an interesting new candidate variant for the development of high myopia, and provide additional data that MMP2 and IL6 have a role in myopia development, but do not support the claim of the title that myopia is caused by an inflammatory reaction.

Thank you for your valuable comments. Four missense mutations in the ZC3H11A gene (c.412G>A, p.V138I; c.128G>A, p.G43E; c.461C>T, p.P154L; and c.2239T>A, p.S747T) were identified in the 1015 HM patients aged from 15 to 18 years. All of the identified mutations exhibited very low frequencies or does not exist in the Genome Aggregation Database (gnomAD) and Clinvar, and using pathogenicity prediction software SIFT, PolyPhen2, and CADD, most of them display high pathogenicity levels. Among them, c.412G>A, c.128G>A and c.461C>T were located in or around a domain named zf-CCCH_3 (Figure 1A and B). Furthermore, all of the mutation sites were located in highly conserved amino acids across different species (Figure 1C). Four mutations resulted in a higher degree of conformational flexibility and altered the negative charge at the corresponding sites (Figure 1D and E). Meanwhile, through transfection of overexpression mutant plasmids, it was found that compared to the wild-type, the mRNA expression levels of IκBα in the nucleus of all four mutant types (ZC3H11AV138I, ZC3H11AG43E, ZC3H11AP154L and ZC3H11AS747T) were significantly reduced (Supplement Figure 3). According to the ACMG guidelines, the above mutations can be classified as “Pathogenic Moderate”.

The existing data show significant differences in refraction between 4 and 10 weeks, with the AL and vitreous cavity depth of Het mice being longer than those of WT mice at 4 and 6 weeks. Although no significant differences were observed at other time points, there was still some degree of increase in these parameters. We continued to measure corneal curvature and found no significant differences between the two groups. Therefore, the difference in refraction may be due to the small size of the mouse eye. A 1 D change in refraction corresponds to only a 5-6 μm change in AL(1). However, the SD-OCT resolution used in this study is relatively low (theoretical resolution of 6 μm)(2, 3), so the small changes measured in vitreous cavity depth and AL may not be statistically significant. Additionally, some studies have shown that axial lengths reported in frozen sections are longer than those measured in vivo for age-matched mice(1, 4). Another possible explanation is that the curvature and refractive power of the lens have changed. These hypotheses provide a reasonable explanation for the mismatch between changes in refraction and ocular length parameters.

To evaluate the change in the number of a specific type of retinal cells, the most commonly used experimental method involves staining with antibodies specific to the target cell type, followed by fluorescence microscopy. The fluorescence intensity or the number of cells can be analyzed semi-quantitatively to assess the changes in the specific cell type in the retina. For example, in retinal degenerative models, rhodopsin-specific staining is used to identify the loss of rod cells. In our study, we selected PCK-α as a marker protein for bipolar cells to assess their number. Additionally, transmission electron microscopy (TEM) was used to observe damage to the cell morphology in the inner nuclear layer (INL) of Het mice, where bipolar cell bodies are located. Based on both sets of data, we conclude that bipolar cells have indeed undergone structural damage and a reduction in number.

Reference

(1) Schmucker C, Schaeffel F. A paraxial schematic eye model for the growing C57BL/6 mouse. Vision research 44, 1857-1867 (2004).

(2) Yuan Y, Chen F, Shen M, Lu F, Wang J. Repeated measurements of the anterior segment during accommodation using long scan depth optical coherence tomography. Eye & contact lens 38, 102-108 (2012).

(3) Shen M, et al. SD-OCT with prolonged scan depth for imaging the anterior segment of the eye. Ophthalmic Surgery, Lasers and Imaging Retina 41, S65-S69 (2010).

(4) Schmucker C, Schaeffel F. In vivo biometry in the mouse eye with low coherence interferometry. Vision research 44, 2445-2456 (2004).

We have removed the hypothesis regarding inflammatory gene changes and myopic development. At present, the explanation is based solely on the correlation of signal pathways, the theoretical basis comes from the reference literature:

“Lin et al., Role of Chronic Inflammation in Myopia Progression: Clinical Evidence and Experimental Validation. EBioMedicine, 2016 Aug:10:269-81, Figure 7.”

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: High myopia (HM) is a severe form of refractive error that results in irreversible visual impairment and even blindness. However, the genetic and pathological mechanisms underlying this condition are not yet fully understood. From a adolescents myopia survey cohort of 1015 HM patients, pathogenic missense mutations were identified in the ZC3H11A gene in four patients by whole exome sequencing. This gene is a zinc finger and stress-induced protein that plays a significant role in regulating nuclear mRNA export. To better understand the function and molecular pathogenesis of myopia in relation to gene mutations, a Zc3h11a knock-out (KO) mouse model was created. The heterozygous KO (Het-KO) mice exhibited significant shifts in vision towards myopia. Electroretinography revealed that the b-wave amplitude was significantly lower in these mice under dark adaptation. Using immunofluorescence antibodies against specific retinal cell types, the density of bipolar cell-labelled proteins was found to be decreased. Transmission electron microscopy findings suggesting ultrastructural abnormalities of the retina and sclera. Retinal transcriptome sequencing showed that 769 genes were differentially expressed, and Zc3h11a was found to have a negative impact on the PI3K-AKT and NF-κB signalling pathways by quantitative PCR and western blotting. In addition, myopia-related factors, including TGF-β1, MMP-2 and IL-6 were found to be upregulated in the retina or sclera. In summary, this study characterized a new pathogenic gene associated with HM. The findings indicated that the ZC3H11A protein may serve as an innate immune and inflammatory response trigger, contributing to the early onset of myopia. These findings offer potential therapeutic intervention targets for controlling the development of HM.

Revised version: High myopia (HM) is a severe form of refractive error that results in irreversible visual impairment and even blindness. However, the genetic and pathological mechanisms underlying this condition are not yet fully understood. From a adolescents myopia survey cohort of 1015 HM patients, pathogenic missense mutations were identified in the ZC3H11A gene in four patients by whole exome sequencing. This gene is a zinc finger and stress-induced protein that plays a significant role in regulating nuclear mRNA export. To better understand the function and molecular pathogenesis of myopia in relation to gene mutations, a Zc3h11a knock-out (KO) mouse model was created. The heterozygous KO (Het-KO) mice exhibited significant shifts in vision towards myopia. Electroretinography revealed that the b-wave amplitude was significantly lower in these mice under dark adaptation. Using immunofluorescence antibodies against specific retinal cell types, the density of bipolar cell-labelled proteins was found to be decreased. Transmission electron microscopy findings suggesting ultrastructural abnormalities of the retina and sclera. Retinal transcriptome sequencing showed that 769 genes were differentially expressed, and Zc3h11a was found to have a negative impact on the PI3K-AKT and NF-κB signaling pathways by quantitative PCR and western blotting. In addition, myopia-related factors, including TGF-β1, MMP-2 and IL-6 were found to be upregulated in the retina or sclera. In summary, this study characterized a new candidate pathogenic gene associated with high myopia, and indicated that the ZC3H11A protein may serve as a stress-induced nuclear response trigger, and its abnormality causes disturbances in a series of inflammatory and myopic factors. These findings offer potential therapeutic intervention targets for controlling the development of HM.

Introduction

Myopia is a highly prevalent eye affliction that commonly develops during childhood and early adolescence. The prevalence of myopia among the adult population ranges from 10-30%, while parts of East and Southeast Asia have reported rates as high as 80-90% among young people (Baird et al., 2020; Holden et al., 2016). HM is a severe refractive error characterised by a diopter ≤ −6.00 D or an ocular axis greater than 26 mm. It is estimated that by 2050, the number of people with HM worldwide will reach 938 million, accounting for 9.8% of the total population. HM can trigger a range of adverse ocular changes, such as cataracts, glaucoma, retinal detachment, macular degeneration and possibly total blindness (Koga et al., 2014; Saw et al., 2005). While conventional methods for the prevention and control of myopia can provide some correction, they are not entirely effective in managing its progression, particularly during childhood and adolescence.

Earlier version: With the development of next-generation sequencing, whole exome sequencing (WES) and whole genome sequencing (WGS) have extended the findings of linkage studies to identify potential causes of syndromic HM (sHM) and non-syndromic HM (nsHM). To date, approximately 20 genes with causal associations have been identified in sHM, including ZNF644, SCO2, CCDC111, LRPAP1, SLC39A5, LEPREL1, P4HA2, OPN1LW, ARR3, BSG, NDUFAF7, CPSF1, TNFRSF21, DZIP1, XYLT1, CTSH, GRM6, LOXL3 and GLRA2 (Haarman et al., 2022; Tian et al., 2023; Yang et al., 2023; Ye et al., 2023). These discoveries have provided insight into the molecular mechanisms underlying HM. However, known candidate genes can only explain about 20% of the causes of this disease (Cai et al., 2019; Tedja et al., 2019). At the same time, the neuromodulators and signal molecules of HM are extremely complex, including sclera extracellular matrix (ECM) remodelling and endoplasmic reticulum (ER) stress (Ikeda et al., 2022), inflammatory responses, the release of dopamine and gamma-aminobutyric acid (GABA), or abnormalities in myopia-related signalling pathways, such as retinoic acid signalling, TGF-β signalling and HIF-1α signalling.

Revised version: With the development of next-generation sequencing, whole exome sequencing (WES) and whole genome sequencing (WGS) have extended the findings of linkage studies to identify potential causes of syndromic HM (sHM) and non-syndromic HM (nsHM). To date, approximately 20 genes with causal associations have been identified in sHM, including ZNF644, SCO2, CCDC111, LRPAP1, SLC39A5, LEPREL1, P4HA2, OPN1LW, ARR3, BSG, NDUFAF7, CPSF1, TNFRSF21, DZIP1, XYLT1, CTSH, GRM6, LOXL3 and GLRA2 (Haarman et al., 2022; Tian et al., 2023; Yang et al., 2023; Ye et al., 2023). These discoveries have provided insight into the molecular mechanisms underlying HM. However, known candidate genes can only explain about 20% of the causes of this disease (Cai et al., 2019; Tedja et al., 2019). At the same time, the neuromodulators and signal molecules of HM are extremely complex, including sclera extracellular matrix (ECM) remodelling and endoplasmic reticulum (ER) stress (Ikeda et al., 2022), inflammatory responses, the release of dopamine and gamma-aminobutyric acid (GABA), or abnormalities in myopia-related signaling pathways, such as retinoic acid signaling, TGF-β signaling and HIF-1α signaling.

Earlier version: Screening mutations from a adolescents myopia survey cohort by WES, we identified zinc finger CCCH domain-containing protein 11A (ZC3H11A) as a HM candidate gene. This particular gene is a member of the zinc finger protein gene family. Several of its members have been linked to myopia or HM, including ZNF644, ZC3H11B, ZFP161 and ZENK. Like ZC3H11B, a gene that is conserved with respect to ZC3H11A in humans, the five genome-wide variants loci have previously been found to be strongly associated with greater axial length (AL) or HM (Schippert et al., 2007; Shi et al., 2011; Szczerkowska et al., 2019; Tang et al., 2020; Wang et al., 2004). Recent proteomic studies have indicated that ZC3H11A may be a constituent of the transcriptional export (TREX) complex. Additionally, it has been suggested that this protein is involved in stress-induced responses (Younis et al., 2018). Dysfunction of ZC3H11A results in enhanced NF-κB signalling through defective IκBα protein expression, which is accompanied by upregulation of numerous innate immune- and inflammation-related mRNAs, including IL-8, IL-6 and TNF in vitro (Jimi et al., 2019). Moreover, patients with myopia have a higher proportion of inflammation-associated cells, including neutrophils, while those with moderate myopia show strained immune system function (Lin et al., 2016; Qi et al., 2022). Conversely, some traditional Chinese medicines can control myopia progression by suppressing AKT and NF-κB mediated inflammatory reactions (Chen et al., 2022). However, the precise pathological mechanism of ZC3H11A in myopia development remains unclear, necessitating further research.

Revised version: Screening mutations from a adolescents myopia survey cohort by WES, we identified zinc finger CCCH domain-containing protein 11A (ZC3H11A) as a HM candidate gene. This particular gene is a member of the zinc finger protein gene family. Several of its members have been linked to myopia or HM, including ZNF644, ZC3H11B, ZFP161 and ZENK. Like ZC3H11B, a gene that is conserved with respect to ZC3H11A in humans, the five genome-wide variants loci have previously been found to be strongly associated with greater axial length (AL) or HM (Schippert et al., 2007; Shi et al., 2011; Szczerkowska et al., 2019; Tang et al., 2020; Wang et al., 2004). Recent proteomic studies have indicated that ZC3H11A may be a constituent of the transcriptional export (TREX) complex. Additionally, it has been suggested that this protein is involved in stress-induced responses (Younis et al., 2018). Dysfunction of ZC3H11A results in enhanced NF-κB signaling through defective IκBα protein expression, which is accompanied by upregulation of numerous innate immune- and inflammation-related mRNAs, including IL-8, IL-6 and TNF in vitro (Jimi et al., 2019). Moreover, patients with myopia have a higher proportion of inflammation-associated cells, including neutrophils, while those with moderate myopia show strained immune system function (Lin et al., 2016; Qi et al., 2022). Conversely, some traditional Chinese medicines can control myopia progression by suppressing AKT and NF-κB mediated inflammatory reactions (Chen et al., 2022). However, the precise pathological mechanism of ZC3H11A in myopia development remains unclear, necessitating further research.

The current study identified four HM-related variants in the ZC3H11A gene among a cohort of 1015 adolescents. Additionally, KO mice were constructed using the CRISPR/Cas9 system and their myopic phenotypes, visual function, bipolar cell apoptosis, and retinal and scleral microstructure were assessed. Moreover, RNA sequencing and expression experiment was performed to identify perturbed molecules and pathways and to examine the interactions between these factors in relation to HM. The above results will provide new ideas for the prevention and control of high myopia, especially early-onset, uncorrectable and familial myopia.

Results

ZC3H11A mutations are associated with HM in a Chinese cohort

Earlier version: The subjects were selected from the Myopia Associated Genetics and Intervention Consortium (MAGIC), our previously published research involving large-scale myopia cohorts and genomic sequencing data. Four missense mutations in the ZC3H11A gene (c.412G>A, p.V138I; c.128G>A, p.G43E; c.461C>T, p.P154L; and c.2239T>A, p.S747T) were identified in the 1015 HM patients aged from 15 to 18 years. The uncorrected visual acuity (UCVA) and ocular axial length of these patients are presented in Table 1. A tessellated fundus was observed in patient 1 and patient 3. All of the identified mutations exhibited very low frequencies in the Exome Aggregation Consortium (ExAC) or Genome Aggregation Database (gnomAD), with CADD > 22.7. Among them, c.412G>A, c.128G>A and c.461C>T were located in or around a domain named zf-CCCH_3 (Figure 1A and B). Furthermore, all of the mutation sites were located in highly conserved amino acids across different species (Figure 1C). Four mutations resulted in a higher degree of conformational flexibility and altered the negative charge at the corresponding sites (Figure 1D and E).

Revised version: The subjects were selected from the Myopia Associated Genetics and Intervention Consortium (MAGIC), which is our previously published study involving large-scale high myopia cohorts and genomic sequencing data. Four missense mutations in the ZC3H11A gene (c.412G>A, p.V138I; c.128G>A, p.G43E; c.461C>T, p.P154L; and c.2239T>A, p.S747T) were identified in the 1015 HM patients aged from 15 to 18 years. The uncorrected visual acuity (UCVA) and axial length of these patients are presented in Table 1. All of the identified mutations exhibited very low frequencies in the Genome Aggregation Database (gnomAD) and Clinvar, and using pathogenicity prediction software SIFT, PolyPhen2, and CADD, most of them display high pathogenicity levels. Among them, c.412G>A, c.128G>A and c.461C>T were located in or around a domain named zf-CCCH_3 (Figure 1A and B). Furthermore, all of the mutation sites were located in highly conserved amino acids across different species (Figure 1C). Four mutations resulted in a higher degree of conformational flexibility and altered the negative charge at the corresponding sites (Figure 1D and E).

Structural modelling of ZC3H11A gene mutations

(A) Illustration of exons and mutation sites of ZC3H11A; exons 5-8 of ZC3H11A encode the zf-CCCH_3 protein domain. (B) A structural model of the full-length human ZC3H11A protein was generated using PyMOL; the locations of the four mutation sites are marked. (C) The p.G43E mutation is located in exon 6, while the p.V138I and p.P154L mutations are in exon 8, and the p.S747T mutation is in exon 20. (D) Multiple sequence alignment of the proteins where ZC3H11A mutations are found in multiple species. (E) These mutations can result in a higher degree of conformational flexibility at the corresponding sites, potentially destabilizing the structural domain.

Zc3h11a Het-KO mice exhibited myopic phenotypes

Earlier version: Het-KO mice were constructed at four weeks of age against a C57BL/6J background using CRISPR/Cas9 technology (Figure supplement 1). Retinal fundus images and ocular histomorphology of Zc3h11a Het-KO mice at eight weeks postnatal were assessed against those of their wild-type (WT) counterparts. No significant structural differences were observed (Figure supplement 2). These findings suggest that the deletion of Zc3h11a does not alter the retinal structure or ocular histomorphology. Therefore, Zc3h11a Het-KO mice are a relevant model for the investigation of the role of Zc3h11a in refractive development.

Revised version: Het-KO mice were constructed at four weeks of age against a C57BL/6J background using CRISPR/Cas9 technology (Supplement Figure 1). Retinal fundus images and ocular histomorphology of Zc3h11a Het-KO mice at eight weeks postnatal were assessed against those of their wild-type (WT) counterparts. No significant structural differences were observed (Supplement Figure 2). These findings suggest that the deletion of Zc3h11a does not alter the retinal structure or ocular histomorphology. Therefore, Zc3h11a Het-KO mice are a relevant model for the investigation of the role of Zc3h11a in refractive development.

Refraction and axial length in Zc3h11a Het-KO mice were found to be significantly greater than in WT littermates (independent samples t-test, p<0.05; Figure 2A and B). The difference between the two genotypes was statistically significant at weeks 4 and 6. Correspondingly, the vitreous chamber depth of Zc3h11a Het-KO mice was deeper than that of WT littermates (independent samples t-test, p<0.05; Figure 2C). There were no significant differences in the anterior chamber depth, lens diameter and body weight between the two groups (Figure 2D-F). These results suggest that the Het-KO of the Zc3h11a gene may lead to an increase in myopia in mice.

Comparison of ocular biometrics and body weights of Zc3h11a Het-KO (n=14) and WT (n=10) littermates during weeks 4-10 of development.

(A, B) Refraction and axial length in Zc3h11a Het-KO mice. (C) Vitreous chamber depth. (D-F) Anterior chamber depth, lens depth and body weight. The effect of genotype on time-dependent refractive development was assessed through independent samples t-tests. P-values are indicated as follows: *P<0.05, **P<0.01, ***P<0.001 and ****<0.0001.

Reduced b-wave amplitude and bipolar cell-labelled protein density in Zc3h11a Het-KO mice

To confirm if Zc3h11a is responsible for refractive development regulation, visual function was assessed by electroretinography (ERG). Upon dark adaptation, b-wave amplitudes in seven-week-old Het-KO mice were significantly lower at dark 3.0 and dark 10.0 compared to WT mice (Figure 3A and C). On the contrary, there were no differences in a-wave amplitudes between the Het-KO and WT groups (Figure 3B). Based on the ERG results, variations between Het-KO and WT littermates can be observed. The ERG b-wave, predominantly generated by bipolar cells, was significantly modified in both the 3.0 and 10.0 dark adaptation conditions. Immunofluorescence analyses were performed on frozen retinal sections to investigate this phenomenon further. Specifically, Zc3h11a, PKCα, Opsin-1 and Rhodopsin markers were utilised to detect the Zc3h11a protein, rod-bipolar cells, cone cells and rod cells, respectively (Figure 3D-F). Additionally, the differences in the Zc3h11a and PKCα protein contents of the retina were quantified through western blot analysis. Notably, the protein densities of Zc3h11a and PKCα were decreased in the Het-KO group (Figure 3G and H). However, there were no significant differences in the densities of cone and rod cells between the two groups (Figure 3I and J).

Electroretinography and immunohistochemical localization of Zc3h11a, PKCα, Opsin-1 and Rhodopsin in WT and Het-KO mice eyes

Earlier version: (A) Representative scotopic ERG responses from Het-KO and WT eyes at dark 0.01, 3.0 and 10.0 cds/m2. (B, C) Statistical analysis of a-wave and b-wave amplitudes. Mean ± standard deviation, n=12/group. (D-F) Immunofluorescence-stained samples to detect Zc3h11a, bipolar cells, cone and rod cells. (G, H) Lower levels of Zc3h11a and PKCα (E. H) expression were observed in Het-KO mice (I, J) The expression levels of Opsin-1 and Rhodopsin between Het-KO and WT mice. Statistical significance was defined as *P<0.05, **P<0.01 and ***P<0.001, as determined by independent samples t-tests.

Revised version: (A) Representative scotopic ERG responses from Het-KO and WT eyes at dark 0.01, 3.0 and 10.0 cds/. (B, C) Statistical analysis of a-wave and b-wave amplitudes. Mean ± standard deviation, n=12/group. (D-F) Immunofluorescence-stained samples to detect Zc3h11a, bipolar cells, cone and rod cells. (G, H) Lower levels of Zc3h11a and PKCα (E. H) expression were observed in Het-KO mice (I, J) the expression levels of Opsin-1 and Rhodopsin between Het-KO and WT mice. Statistical significance was defined as *P<0.05, **P<0.01 and ***P<0.001, as determined by independent samples t-tests.

Retinal ultrastructure alterations in het-KO mice model

Transmission electron microscopy (TEM) was used to analyse retinal ultrastructure of in 10-week-old mice. In the inner nuclear layer (INL) of the retina, in contrast to WT mice, Het-KO mouse cells had enlarged perinuclear gaps (black arrow), perinuclear cytoplasmic oedema (blue arrow), and thinned and lightened cytoplasm (Figure 4A and B). However, in the outer nuclear layer (ONL), no significant difference in cell morphology was observed between the two groups of mice (Figure 4C and D). Furthermore, when compared to the WT group, the Het-KO mice exhibited relatively damaged photoreceptor cell membrane discs (MB), in which the outer layer was detached and sparsely distributed locally (Figure 4E and F). There were also a small number of broken membrane discs, as well as some disorganized and loosely arranged ones (red arrow), with a slight increase in size (Figure 4G and H).

TEM showing the retinal ultrastructure of WT and Het-KO mice.

(A, B) Cellular morphology of the INL in which bipolar cells are located in WT (A) and Het-KO mice (B). (C, D) Cellular morphology of the ONL in which the optic cells are located in WT (C) and Het-KO mice (D). (E, F) MB structure in WT mice. (G, H) MB structure of Het-KO mice.

RNA-Seq analysis of molecular and pathways changes in Zc3h11a Het-KO mice retinas

Earlier version: In the retina transcriptome analysis, 769 genes were differentially expressed (Fold change (FC) of at least two and a P value < 0.05) in the Zc3h11a Het-KO group, of which, 303 were upregulated and 466 were downregulated (Figure 5A and B). GO enrichment analysis revealed that these were primarily involved in biological processes (such as zinc ion transmembrane transport, RNA biosynthetic and metabolic process, and negative regulation of NIK/NF-κb signalling) and molecular functions (such as calcium ion binding and zinc ion transmembrane transporter activity) (Figure 5C and D). KEGG pathway enrichment analysis indicated that the differentially expression genes (DEGs) in the Zc3h11a Het-KO group were primarily involved in the PI3K-AKT signalling pathway, MAPK signalling pathway, human virus infection and neuropsychiatric disease (Figure 5E).

Revised version: In the retina transcriptome analysis, 769 genes were differentially expressed (Fold change (FC) of at least two and a P value < 0.05) in the Zc3h11a Het-KO group, of which, 303 were upregulated and 466 were downregulated (Figure 5A and B). GO enrichment analysis revealed that these were primarily involved in biological processes (such as zinc ion transmembrane transport, RNA biosynthetic and metabolic process, and negative regulation of NIK/NF-κb signaling) and molecular functions (such as calcium ion binding and zinc ion transmembrane transporter activity) (Figure 5C and D). KEGG pathway enrichment analysis indicated that the differentially expression genes (DEGs) in the Zc3h11a Het-KO group were primarily involved in the PI3K-AKT signaling pathway, MAPK signaling pathway, and neuropsychiatric disease (Figure 5E).

RNA-Seq analysis of molecular and pathway differences between WT control and Het-KO retinas.

Earlier version: (A, B) Volcano maps were utilized to visualize the differentially expressed genes (DEGs) between WT control and Het-KO mouse retinas. (C, D) GO enrichment analysis of DEGs in the Het-KO group (E) KEGG pathway enrichment analysis of DEGs in the Het-KO group.

Revised version: (A) Volcano maps were utilized to visualize the differentially expressed genes (DEGs) between WT control and Het-KO mouse retinas. (B, C) GO enrichment analysis of DEGs in the Het-KO group. (D) KEGG pathway enrichment analysis of DEGs in the Het-KO group.

Zc3h11a negatively regulates the PI3K-AKT and NF-κB pathways

Earlier version: PI3K-AKT is one of the most important pathways for cell survival, division, autophagy and differentiation (Alzahrani, 2019). Meanwhile, it has been found that ZC3H11A can regulate the NF-κB pathway at the level of IκB mRNA export in ZC3-KO cells (Darweesh et al., 2022). IκBα is a critical protein that governs NF-κB function predominantly within the cytoplasmic compartment. It plays a pivotal role in inhibiting the activation of NF-Κb (Dyson and Komives, 2012; Manavalan et al., 2010). In addition, AKT promotes IκBα phosphorylation to undergo degradation, which enhances the nuclear translocation of NF-kB (Alzahrani, 2019; Torrealba et al., 2020). Accordingly, we investigated whether ZC3H11A has any regulatory influence on the transcriptional activation of PI3K-AKT and NF-κB in vivo. To this end, we evaluated the mRNA and protein expression levels of ZC3H11A, PI3K, AKT, p-AKT, IκBα and NF-κB in the retina. The expression levels of ZC3H11A, PI3K, AKT, p-AKT and NF-κB were significantly upregulated, while IκBα was downregulated in Het-KO mice (Figure 6A-L). These findings strongly suggest that ZC3H11A exerts a negative impact on the PI3K-AKT and NF-κB pathways.

Revised version: PI3K-AKT is one of the most important pathways for cell survival, division, autophagy and differentiation (Alzahrani, 2019). Meanwhile, it has been found that ZC3H11A can regulate the NF-κB pathway at the level of IκB mRNA export in ZC3-KO cells (Darweesh et al., 2022). IκBα is a critical protein that governs NF-κB function predominantly within the cytoplasmic compartment. It plays a pivotal role in inhibiting the activation of NF-Κb (Dyson and Komives, 2012; Manavalan et al., 2010). In addition, AKT promotes IκBα phosphorylation to undergo degradation, which enhances the nuclear translocation of NF-kB (Alzahrani, 2019; Torrealba et al., 2020). Accordingly, we investigated whether ZC3H11A has any regulatory influence on the transcriptional activation of PI3K-AKT and NF-κB in vivo. To this end, we evaluated the mRNA and protein expression levels of ZC3H11A, PI3K, AKT, p-AKT, IκBα and NF-κB in the retina. The expression levels of ZC3H11A, PI3K, AKT, p-AKT and NF-κB were significantly upregulated, while IκBα was downregulated in Het-KO mice (Figure 6A-L). Meanwhile, through transfection of overexpression mutant plasmids, it was found that compared to the wild-type, the mRNA expression levels of IκBα in the nucleus of all four mutant types (ZC3H11AV138I, ZC3H11AG43E, ZC3H11AP154L and ZC3H11AS747T) were significantly reduced (Supplement Figure 3). These findings strongly suggest that ZC3H11A exerts a negative impact on the PI3K-AKT and NF-κB pathways.

Earlier version: Effect of Zc3h11a on the expression of genes and proteins related to the PI3K-AKT and NF-κB signalling pathways in the retina.

Revised version: Effect of Zc3h11a on the expression of genes and proteins related to the PI3K-AKT and NF-κB signaling pathways in the retina

(A-C) The ZC3H11A, PI3K and AKT gene levels (n=3) were evaluated. (D, E) The IκBα and NF-κB gene levels (n=3). (F) Protein expression levels of ZC3H11A, PI3K, AKT, p-AKT and GAPDH were detected by western blot. (G-I) Quantitative analysis of the ZC3H11A and PI3K levels were normalized to GAPDH, while p-AKT levels were normalized to AKT (n=3). (J) The protein expression levels of IκBα, NF-κB and GAPDH were determined by western blot analysis. (K, L) Quantitative analyses of IκBα and NF-κB were normalized to GAPDH (n=3).

TGF-β1, MMP-2 and IL-6 were increased in Zc3h11a Het-KO mice

Earlier version: In animal models of myopia, it has been established that intraocular concentrations of TGF-β1 are elevated (Chen et al., 2013; Liu et al., 2022), especially in the retina and sclera. The abnormal expression of matrix metalloproteinase-2 (MMP-2) and interleukin-6 (IL-6) can induce myopia, and these are regulated by NF-κB (Libermann and Baltimore, 1990; Wu and Schmid-Schönbein, 2011). To investigate the hypothesis that the TGF-β1 and NF-κB signalling pathways exhibit hyperactivity in the retina in the presence of reduced Zc3h11a, we assessed the expression of TGF-β1, MMP-2 and IL-6 at the mRNA levels in the retina and sclera of WT and Het-KO mice, while analysing scleral TEM. Additionally, we confirmed changes in the expression levels of TGF-β1 and MMP-2 in the retina at the protein level. qPCR analysis showed increased expression of TGF-β1, MMP-2, and IL-6 in the sclera and retina of Het-KO mice compared with WT mice (Figure 7A and B). Scleral TEM results suggested that scleral collagen fibres in Het-KO mice were disorganized over a large area with irregular transverse and longitudinal arrangement (Figure 7C). The western blot results showed increased expression levels of TGF-β1 and MMP-2 in the retinas of Zc3h11a Het-KO mice (Figure 7D).

Revised version: In animal models of myopia, it has been established that intraocular concentrations of TGF-β1 are elevated (Chen et al., 2013; Liu et al., 2022), especially in the retina and sclera. The abnormal expression of matrix metalloproteinase-2 (MMP-2) and interleukin-6 (IL-6) can induce myopia, and these are regulated by NF-κB (Libermann and Baltimore, 1990; Wu and Schmid-Schönbein, 2011). To investigate the hypothesis that the TGF-β1 and NF-κB signaling pathways exhibit hyperactivity in the retina in the presence of reduced Zc3h11a, we assessed the expression of TGF-β1, MMP-2 and IL-6 at the mRNA levels in the retina and sclera of WT and Het-KO mice, while analysing scleral TEM. Additionally, we confirmed changes in the expression levels of TGF-β1 and MMP-2 in the retina at the protein level. qPCR analysis showed increased expression of TGF-β1, MMP-2, and IL-6 in the sclera and retina of Het-KO mice compared with WT mice (Figure 7A and B). Scleral TEM results suggested that scleral collagen fibres in Het-KO mice were disorganized over a large area with irregular transverse and longitudinal arrangement (Figure 7C). The western blot results showed increased expression levels of TGF-β1 and MMP-2 in the retinas of Zc3h11a Het-KO mice (Figure 7D).

Effect of Zc3h11a on the expression of TGF-β1, MMP-2 and IL-6 in the retina and sclera

(A, B) Examination of TGF-β1, MMP-2 and IL-6 gene expression in the sclera (A) and retina (B) (n=3). (C) TEM structure of sclera. (D) Determination of TGF-β1 and MMP-2 protein expression levels, normalized to GAPDH, through western blot analysis (n=3).

Discussion

Removed in the revised version: This study identified and validated a new candidate gene in a large high myopia cohort, ZC3H11A. Moreover, the study provided evidence of the presence of moderate or HM phenotypes and damaged bipolar cells in vivo by constructing Het-KO mice for the first time. TEM analysis showed ultrastructural changes in parts of both the retina and sclera of Het-KO mice. To understand how ZC3H11A regulates the development of HM, retinal transcriptome sequencing was performed. The results revealed that Zc3h11a upregulates the PI3K-AKT and NF-κB signalling pathways. In addition, the expression levels of myopia-related factors, such as TGF-β1, MMP-2 and IL-6, were upregulated in the retina and sclera of the Het-KO mice. Therefore, it can be speculated that the ZC3H11A protein may act as an innate immune and inflammatory response trigger, contributing to the early onset of HM.

Earlier version: ZC3H11A is involved in the export and post-transcriptional regulation of selected mRNA transcripts required to maintain metabolic processes in embryonic cells; these are essential for the viability of early mouse embryos (Younis et al., 2023). In studies of chickens, mice and humans, the zinc finger protein ZENK was found to play a role in the development of refractive myopic excursion and lengthening of the ocular axis. Moreover, early growth response gene type 1 Egr-1 (the human homolog of ZENK) activates the TGF-β1 gene by binding to its promoter, which is thought to be associated with myopia (Baron et al., 2006; Xiao et al., 2022). Another zinc protein finger protein 644 isoform, ZNF644, has recently been identified by WES as causing HM in Han Chinese families (Shi et al., 2011). ZC3H11A is also a zinc finger protein, and the findings of the current study revealed that Zc3h11a Het-KO mice showed a myopic shift, which is consistent with previous studies reporting reduced protein expression of Zc3h11a in an unilateral induced myopic mouse model (Fan et al., 2012).

Revised version: This study identified and validated a new candidate gene in a large high myopia cohort, ZC3H11A. Moreover, the study provided evidence of the presence of moderate or HM phenotypes and damaged bipolar cells in vivo by constructing Het-KO mice for the first time. TEM analysis showed ultrastructural changes in parts of both the retina and sclera of Het-KO mice. To understand how ZC3H11A regulates the development of HM, retinal transcriptome sequencing was performed. The results revealed that Zc3h11a upregulates the PI3K-AKT and NF-κB signaling pathways. In addition, the expression levels of myopia-related factors, such as TGF-β1, MMP-2 and IL-6, were upregulated in the retina and sclera of the Het-KO mice. Therefore, it can be speculated that the ZC3H11A protein may act as a stress-induced nuclear response trigger, its abnormal expression contributing to the early onset of HM. ZC3H11A is involved in the export and post-transcriptional regulation of selected mRNA transcripts required to maintain metabolic processes in embryonic cells, these are essential for the viability of early mouse embryos (Younis et al., 2023). In studies of chickens, mice and humans, the zinc finger protein ZENK was found to play a role in the development of refractive myopic excursion and lengthening of the ocular axis. Moreover, early growth response gene type 1 Egr-1 (the human homolog of ZENK) activates the TGF-β1 gene by binding to its promoter, which is thought to be associated with myopia (Baron et al., 2006; Xiao et al., 2022). Another zinc protein finger protein 644 isoform, ZNF644, has recently been identified by WES as causing HM in Han Chinese families (Shi et al., 2011). ZC3H11A is also a zinc finger protein, and the findings of the current study revealed that Zc3h11a Het-KO mice showed a myopic shift, which is consistent with previous studies reporting reduced protein expression of Zc3h11a in an unilateral induced myopic mouse model (Fan et al., 2012).

Earlier version: In vertebrate models, refractive development and ocular axial growth are visually controlled (Wallman and Winawer, 2004). The regulation of axial length or refractive error occurs through complex light-dependent retina-to-sclera signalling (Tkatchenko and Tkatchenko, 2019). Optical scatter information is processed by the retina and then converted into molecular signals that regulate peripheral retinal growth and scleral connective tissue renewal, ultimately affecting the growth rate of the posterior segment of the eye (Harper and Summers, 2015; Tkatchenko et al., 2006; Tkatchenko et al., 2018). All cell types of the retina contain myopia-related genes and retinal circuitry driving refractive error. In this study, Het-KO mice were found to have reduced b-wave amplitudes, diminished relative fluorescence intensity of bipolar cells and ultrastructural changes in inner nuclear layer (INL) of the retina. Therefore, it is presumed that retinal circuitry changes result in impairment of visual signal processing and transduction causing HM.

Revised version: In vertebrate models, refractive development and ocular axial growth are visually controlled (Wallman and Winawer, 2004). The regulation of axial length or refractive error occurs through complex light-dependent retina-to-sclera signaling (Tkatchenko and Tkatchenko, 2019). Optical scatter information is processed by the retina and then converted into molecular signals that regulate peripheral retinal growth and scleral connective tissue renewal, ultimately affecting the growth rate of the posterior segment of the eye (Harper and Summers, 2015; Tkatchenko et al., 2006; Tkatchenko et al., 2018). All cell types of the retina contain myopia-related genes and retinal circuitry driving refractive error. In this study, Het-KO mice were found to have reduced b-wave amplitudes, diminished relative fluorescence intensity of bipolar cells and ultrastructural changes in inner nuclear layer (INL) of the retina. Therefore, it is presumed that retinal circuitry changes result in impairment of visual signal processing and transduction causing HM.

Earlier version: Aberrant activation of the PI3K-AKT and NF-κB signalling pathways has been identified in highly myopic retinas. Following the sequencing evidence indicating elevations in PI3K-AKT and NF-κB signalling (Lin et al., 2016), we sought to explore the potential interplay between these pathways in mice. The results revealed that a decrease in Zc3h11a inhibited the translocation of IκBα from the nucleus to the cytoplasm. IκBα is a downstream factor of PI3K-AKT and binds to NF-κB (p65) in the cytoplasm, inhibiting its nuclear translocation. The qPCR and western Blot results verified that Zc3h11a negatively regulates both the PI3K-AKT and NF-κB signalling pathways. Thus, Zc3h11a may exert an influence on these pathways by modulating the cytoplasmic levels of IκBα, contributing to the development of myopia.

Revised version: Aberrant activation of the PI3K-AKT and NF-κB signaling pathways has been identified in highly myopic retinas. Following the sequencing evidence indicating elevations in PI3K-AKT and NF-κB signaling (Lin et al., 2016), we sought to explore the potential interplay between these pathways in mice. The results revealed that a decrease in Zc3h11a inhibited the translocation of IκBα from the nucleus to the cytoplasm. IκBα is a downstream factor of PI3K-AKT and binds to NF-κB (p65) in the cytoplasm, inhibiting its nuclear translocation. The qPCR and western Blot results verified that Zc3h11a negatively regulates both the PI3K-AKT and NF-κB signaling pathways. Thus, Zc3h11a may exert an influence on these pathways by modulating the cytoplasmic levels of IκBα, contributing to the development of myopia.

Earlier version: NF-κB was first discovered 25 years ago and described as a key regulator of induced gene expression in the immune system, playing a central role in the coordinated control of intrinsic immune and inflammatory responses (Hayden and Ghosh, 2011; Morgan and Liu, 2011). The downstream factors of this pathway include IL-6, IL-8, TNF-α, MMP-2, TGF-β1, etc (Jimi et al., 2019; Yoshida and Whitsett, 2006). Studies have demonstrated that the MMP-2 and IL-6 expression levels are increased in myopic eyes and that inhibiting MMP-2 or IL-6 expression will provide some degree of control over myopia progression (Lin et al., 2016; Zhao et al., 2018). Identification of the PI3K-AKT-NF-κB signalling pathway and downstream myopia effect in our study may open up new opportunities for the prevention and treatment of HM and fundus lesions in the future.

Revised version: NF-κB was first discovered 25 years ago and described as a key regulator of induced gene expression in the immune system, playing a central role in the coordinated control of intrinsic immune and inflammatory responses (Hayden and Ghosh, 2011; Morgan and Liu, 2011). The downstream factors of this pathway include IL-6, IL-8, TNF-α, MMP-2, TGF-β1, etc (Jimi et al., 2019; Yoshida and Whitsett, 2006). Studies have demonstrated that the MMP-2 and IL-6 expression levels are increased in myopic eyes and that inhibiting MMP-2 or IL-6 expression will provide some degree of control over myopia progression (Lin et al., 2016; Zhao et al., 2018). Identification of the PI3K-AKT-NF-κB signaling pathway and downstream myopia effect in our study may open up new opportunities for the prevention and treatment of HM and fundus lesions in the future.

Earlier version: There are a number of limitations of this study that should be acknowledged. First, Zc3h11a homozygous KO (Homo-KO) mice were not obtained in our study because homozygous deletion of exons may confer embryonic lethality (Younis et al., 2023). Second, while the current study observed a reduced b-wave amplitude and bipolar cell fluorescence staining intensity in Het-KO mice, an in-depth exploration of the underlying mechanism was beyond the scope of the research. Finally, there was a limited observation time in this study and the effect of ZC3H11A on the late refractive system was not assessed.

Revised version: There are a number of limitations of this study that should be acknowledged. First, the patient’s ophthalmic examination was not comprehensive enough, and genetic analysis such as segregation analysis was not performed. Second, Zc3h11a homozygous KO (Homo-KO) mice were not obtained in our study because homozygous deletion of exons may confer embryonic lethality (Younis et al., 2023). Third, while the current study observed a reduced b-wave amplitude and bipolar cell fluorescence staining intensity in Het-KO mice, an in-depth exploration of the underlying mechanism was beyond the scope of the research. Finally, there was a limited observation time in this study and the effect of ZC3H11A on the late refractive system was not assessed.

Earlier version: Overall, this study has provided four key findings. First, it was confirmed that the ZC3H11A is a new candidate gene is associated with HM in a cohort of Chinese Han individuals. Secondly, changes in Zc3h11a expression were found to affect retinal function, particularly in bipolar cells. Third, changes in Zc3h11a expression were found to cause alterations in numerous signalling pathways, the most notable being the PI3K-AKT and NF-κB pathways. Fourth, changes in Zc3h11a expression were found to cause changes in the myopia-related genes TGF-β1, MMP-2 and IL-6. The above results suggest that Zc3h11a promotes the translocation of IκBα from the nucleus to the cytoplasm, thereby exerting negative feedback regulation of the PI3K-AKT signalling pathway and inhibiting the NF-κB signalling pathway. At the same time, the increase in TGF-β1 in myopic eyes stimulates the PI3K-AKT signalling pathway, which leads to activation of the PI3K-AKT signalling pathway, resulting in downstream degradation of IκBα after phosphorylation. This, in turn, increases the nuclear translocation of NF-κB. To sum up, Zc3h11a promotes the myopia-associated factor TGF-β1 by acting directly or indirectly on both the PI3K-AKT and NF-κB signalling pathway-mediated inflammatory reactions and the expression of MMP-2 and IL-6, which together, promote the development of early myopia (Figure 8).

Revised version: Overall, this study has provided four key findings. First, it was confirmed that the ZC3H11A is a new candidate gene associated with HM in a cohort of Chinese Han individuals. Secondly, changes in Zc3h11a expression were found to affect retinal function, particularly in bipolar cells. Third, changes in Zc3h11a expression were found to cause alterations in numerous signaling pathways, the most notable being the PI3K-AKT and NF-κB pathways. Fourth, changes in Zc3h11a expression were found to cause changes in the myopia-related genes TGF-β1, MMP-2 and IL-6. The above results suggest that abnormal expression of Zc3h11a triggers nuclear translocation defects of IκBα, thereby exerting negative feedback regulation of the PI3K-AKT signaling pathway and inhibiting the NF-κB signaling pathway. At the same time, the increase in TGF-β1 in myopic eyes stimulates the PI3K-AKT signaling pathway, which leads to activation of the PI3K-AKT signaling pathway, resulting in downstream degradation of IκBα after phosphorylation. This, in turn, increases the nuclear translocation of NF-κB. To sum up, Zc3h11a promotes the myopia-associated factor TGF-β1 by acting directly or indirectly on both the PI3K-AKT and NF-κB signaling pathway-mediated stress reactions and the expression of MMP-2 and IL-6, which together, promote the development of early high myopia (Figure 8). Therefore, this model can be used as one of the new strategies for intervention and treatment of high myopia. However, because the causes of myopia or HM are complex and involve different tissues and molecular pathways in the eye, it is likely that future research will identify more genes and molecular mechanisms that could provide guidance for clinical intervention and treatment.

Earlier version: Schematic representation of the PI3K-AKT-NF-κB signalling pathway in retina.

Revised version: Schematic representation of the PI3K-AKT-NF-κB signaling pathway in retina

Removed in the revised version: Therefore, this model can be used as one of the new strategies for intervention and treatment of high myopia. However, because the causes of myopia or HM are complex and involve different tissues and molecular pathways in the eye, it is likely that future research will identify more genes and molecular mechanisms that could provide guidance for clinical intervention and treatment.

Materials and methods

Ethics approval and consent to participate

The studies that involved human participants were approved by the Eye Hospital, Wenzhou Medical University (Wenzhou, China), and were carried out in strict adherence to the guidelines of the Helsinki Declaration. All participants provided informed consent. Furthermore, all animal experiments were approved by the Animal Care and Ethics Committee at Wenzhou Medical University and were performed according to the guidelines set forth in the Association for Research in Vision and Ophthalmology Statement for the Use of Animals in Ophthalmic and Visual Research. (https://www.arvo.org/About/policies/statement-for-the-use-of-animals-in-ophthalmic-and-vision-research/).

Recruitment of subjects

Earlier version: A total of 1015 sporadic HM patients were recruited from the Eye Hospital of Wenzhou Medical University for MAGIC project. HM was defined as SE ≤ −6.00 D in either eye (Flitcroft et al., 2019). Detailed ophthalmic examinations, including visual acuity, axial length, slit-lamp, fundus photography and optical coherence tomography were performed.

Revised version: A total of 1015 sporadic HM patients (aged from 15-18) were recruited from the Eye Hospital of Wenzhou Medical University for MAGIC project (Su et al., 2023). HM was defined as SE ≤ −6.00 D in either eye (Flitcroft et al., 2019). Some patients underwent certain ophthalmic examinations, including visual acuity (at least three measurements under non-cycloplegia), axial length and fundus photography were performed.

WES and the detection of variants

DNA was extracted from all probands using a FlexiGene DNA Kit (Qiagen, Venlo, The Netherlands), according to the manufacturer’s protocol. DNA from all probands underwent WES using a Twist Human Core Exome Kit and an Illumina NovaSeq 6000 sequencing system (150PE) (Berry Genomics Institute, Beijing, China). The mean depth and coverage of the target region were approximately 78.26× and 99.7%, respectively. Sequence reads were aligned to the reference human genome (UCSC hg19) using the Burrows-Wheeler aligner (BWA). The variants were called and annotated with Verita Trekker and Enliven software (Transcript ID, ENST00000332127.4, Berrry Genomics Institute), respectively.

Het-KO mice model construction, husbandry and ocular biometric measurements

Het-KO mice were constructed at four weeks of age against a C57BL/6J background using CRISPR/Cas9 technology at GemPharmatech Co.,Ltd (Nanjing, China). Exon5-exon6 of the Zc3h11a transcript was recommended as the KO region; this region contains a 244bp coding sequence. Littermates of Het-KO and WT mice were used for all experiments. All animals were housed in the animal husbandry room of Wenzhou Medical University. Mice were housed in standard transparent mouse cages at 22±2°C with a 12-hour light/12-hour dark cycle (light from 8 am to 8 pm, brightness approximately 200-300 lux) and free access to food and water. To exclude the effect of ocular developmental malformations, mice with small eyes and ocular lesions were excluded from the observation cohort. Body weight, refraction and ocular biometrics of Het-KO (n=14) and WT (n=10) mice were assessed at 4, 5, 6, 8 and 10 weeks of age.

Electroretinography

To evaluate the effect of Zc3h11a on the electrophysiological properties of various neuronal populations in the retina, ERG was performed on the right eye of Het-KO and WT littermates at seven weeks of age, at the same time of day (n=12). Both scotopic and photopic ERG responses were evaluated. The mice were dark-adapted overnight, and all procedures were performed under dim red light (<1 lux). The animals were anaesthetized with intraperitoneal injection of pentobarbital sodium (40 mg/kg) and the pupils were dilated with 0.5% tropicamide. A heating table (37[) was used to maintain body temperature. ERG was recorded using a Roland Electrophysiological System (RETI-Port21, Roland Consult, Germany) with ring-shaped corneal electrodes. Three light intensities (0.01, 3.0 and 10.0 cds/m2) were applied to each animal. The following parameters were measured: a-wave amplitude and b-wave amplitude. Photopic ERG was measured at a light intensity of 3.0 cds/m2 after 10 min light adaption at 25 cds/m2. The responses of the Het-KO eyes (right) were compared with those of the WT eyes (right).

Immunofluorescence

Whole mouse eyes (10 weeks, littermates Het-KO and WT) were fixed with 4% para-formaldehyde (BL539A, Biosharp, China) for 1 h and 30% sucrose for dehydration overnight. Then, they were fabric-embedded in a frozen fabric matrix compound at - 20[. Prepared tissue blocks were segmented with a cryostat at a thickness of 12 microns and collected on clean adhesive slides. The slices containing the sections were air-dried at room temperature (RT) for 15 min. After washing 3× with 1× PBS for 5 min, 5% BSA (SW3015, Solarbio, China) and 0.03% triton-X100 (P0096, Beyotime, China) diluted with 1× PBS were added as permeable membrane-blocking buffers. The slides were incubated for 1 h at RT in a humid chamber. Then, they were incubated overnight at 4[ with the specific primary antibody Zc3h11a (1:50, 26081, Proteintech, USA). PKC α (1:200, ab32518, Abcam, UK), Opsin-1 (1:200, NB110-74730, Novus Biologicals, USA) and Rhodopsin (1:200, NBP2-25160, Novus Biologicals, USA) were added and the slices were incubated further at 4[ in a humid chamber overnight. After washing 4× with 1× PBS for 6 min, goat anti-rabbit Alexa Fluor 488 (1:500, ab150077, Abcam, UK) was added and the slides were incubated for 90 min at RT. After washing, the slides were mounted with an antifade medium containing DAPI (P0131, Beyotime, China) to visualize the cell nucleus. Sections incubated with 5% BSA and without primary antibodies were used as negative controls. A fluorescence microscope (LSM 880, ZEISS, Germany) was used to examine the slides and capture images. The experiments were repeated in duplicate with three different samples.

Transmission electron microscopy

At 10 weeks of age, two mice of each genotype were euthanized, and their eyes were removed. The eyes were fixed in a solution containing 2.5% glutaraldehyde and 0.01 M phosphate buffer (PB) (pH 7.0-7.5) for 15 min while the optic cups were dissected. The optic cups were then fixed with 1% osmium acid at room temperature away from light for 2h, after which they were rinsed three times with 0.1 MPB for 15min each time. Tissues were sequentially dehydrated in 30%-50%-70%-80%-95%-100%-100% ethanol upstream for 20 min each time, and 100% acetone twice for 15 min each time. Finally, the tissues were embedded in epoxy resin (Polybed 812) mixed 1:1 with acetone. Ultrathin sections were prepared using diamond knives and an EM UC7 ultramicrotome (Leica, Germany), then the sections were stained with 2% aqueous dioxygen acetate and 1% phosphotungstic acid (pH 3.2). Finally, the structures were examined using a transmission electron microscope (HT7700, Hitachi, Tokyo, Japan).

Added in the revised version: Cell culture, plasmid construction and transfected

Added in the revised version: HEK293T cell line (from Chinese Academy of Sciences) was cultured in DMEM medium containing 10% high-quality 10% fetal bovine serum (FBS) and 1% penicillin streptomycin dual antibody. Using restriction endonucleases HindIII and EcoRI to cleave the target gene product (wild type, ZC3H11AV138I, ZC3H11AG43E, ZC3H11AP154L and ZC3H11AS747T) and vector pcDNA3.1 (+), and purify the fragments. Then, under the action of T4 DNA ligase, the target gene was ligated with the vector pcDNA3.1 (+), and the ligated product was transformed into competent bacterial cells. Sequencing was used to identify and select the successfully constructed target gene expression plasmid vector. Finally, a sufficient amount of vector plasmids were obtained through ultrapure endotoxin extraction. Transfected with wild type, ZC3H11AV138I, ZC3H11AG43E, ZC3H11AP154L and ZC3H11AS747T overexpression plasmids using NDE3000 reagent (Western, China).

RNA sequencing analysis of molecular and pathway changes in the mouse retina

Mice aged four weeks were executed by cervical dislocation after CO2 asphyxiation. Then, the eyes were enucleated and dissected to obtain the entire retinas. Three eye retinas from the Zc3h11a Het-KO and WT groups were used for the mRNA sequencing experiments and performed by Biomarker Biotechnology Co. (Beijing, China). The mRNA sequences were mapped to the genome (GRCm38). WebGestalt (http://www.webgestalt.org/) was used to generate GO terms and KEGG pathways.

Earlier version: qRT-PCR

Revised version: RNA extraction and qRT-PCR

Earlier version: To determine the reliability of the transcriptome results, qPCR validation was performed. RNA was isolated using Trizol reagent (RC112, Vazyme, China), and the purity was confirmed by the OD260/280 nm absorption ratio (1.9-2.1) (Nanodrop 2000, Thermo Scientific, USA). Total RNA (2 µg) was reverse transcribed to cDNA using a cDNA Synthesis Kit (R323, Vazyme, China). qPCR was performed with a RT-PCR detection system (Applied Biosystems, California, USA) using a SYBR Premix Ex Taq Kit (Q711, Vazyme, China), according to the manufacturer’s instructions. qRT-PCR was performed (ABI-Q6, California, USA) in a 20 µL reaction, under the following conditions: 95°C for 10 min, followed by 40 cycles of amplification at 95°C for 10 s and 60°C for 60 s. Melting curve analysis was used to determine specific amplification. All experiments were performed in triplicate. Relative quantification was performed using the ΔΔCt method. The specific gene products were amplified using the following primer pairs (Table supplement 1).

Revised version: To determine the reliability of the transcriptome results, qPCR validation was performed. Total RNA and nuclear RNA were extracted using Trizol reagent (RC112, Vazyme, China) and Cytodynamic & Nuclear RNA Purification Kit (21000, Norgen Biotek, Canada), respectively, and the purity was confirmed by the OD260/280 nm absorption ratio (1.9-2.1) (Nanodrop 2000, Thermo Scientific, USA). Total RNA (2 µg) was reverse transcribed to cDNA using a cDNA Synthesis Kit (R323, Vazyme, China). qPCR was performed with a RT-PCR detection system (Applied Biosystems, California, USA) using a SYBR Premix Ex Taq Kit (Q711, Vazyme, China), according to the manufacturer’s instructions. qRT-PCR was performed (ABI-Q6, California, USA) in a 20 µL reaction, under the following conditions: 95°C for 10 min, followed by 40 cycles of amplification at 95°C for 10 s and 60°C for 60 s. Melting curve analysis was used to determine specific amplification. All experiments were performed in triplicate. Relative quantification was performed using the ΔΔCt method. The specific gene products were amplified using the following primer pairs (Supplement Table 1).

Western blot

Earlier version: The mouse retina was separated immediately after enucleation of the eyeball and lysed in RIPA (10 mM Tris-Cl, 100 mM NaCl, 1 mM EDTA, 1 mM NaF, 20 mM Na4P2O7, 2 mM Na3VO4, 1% Triton X-100, 10% glycerol, 0.1% sodium dodecyl sulphate and 0.5% deoxycholate) lysis buffer containing protease and phosphatase inhibitors (P1045, Beyotime, China). Equal amounts of protein (15 μg) were separated on 10% Tris-glycine gel, transferred to PVDF (polyvinylidene fluoride) membranes, and blocked with 5% skim milk. The primary antibodies used included ZC3H11A (ab99930, Abcam, UK), AKT (#4691, Cell Signaling Technology, USA), p-AKT (#4060, Cell Signaling Technology, USA), IκBα (ab32518, Abcam, UK), NF-κB (ab32536, Abcam, UK), TGF-β1 (21898, Proteintech, USA) and MMP-2 (ab92536, Abcam, UK) diluted to 1:1000 in TBST-5% milk. Then, the membranes were incubated with goat anti-rabbit IgG conjugated with HRP (1:2000 in TBST-5% milk) (SA00001-2, Proteintech, USA) for 90 min at room temperature and developed using western blotting reagents (BL523B, Bioshark, China). GAPDH (AF2823, Byotime, China) was used as the internal control.

Revised version: The mouse retina was separated immediately after enucleation of the eyeball and lysed in RIPA (10 mM Tris-Cl, 100 mM NaCl, 1 mM EDTA, 1 mM NaF, 20 mM Na4P207, 2 mM Na3 V04, 1% Triton X-100, 10% glycerol, 0.1% sodium dodecyl sulphate and 0.5% deoxycholate) lysis buffer containing protease and phosphatase inhibitors (P1045, Beyotime, China). Equal amounts of protein (15 μg) were separated on 10% Tris-glycine gel, transferred to PVDF (polyvinylidene fluoride) membranes, and blocked with 5% skim milk. The primary antibodies used included ZC3H11A (ab99930, Abcam, UK), AKT (#4691, Cell Signaling Technology, USA), p-AKT (#4060, Cell Signaling Technology, USA), IκBα (ab32518, Abcam, UK), NF-κB (ab32536, Abcam, UK), TGF-β1 (21898, Proteintech, USA) and MMP-2 (ab92536, Abcam, UK) diluted to 1:1000 in TBST-5% milk. Then, the membranes were incubated with goat anti-rabbit IgG conjugated with HRP (1:2000 in TBST-5% milk) (SA00001-2, Proteintech, USA) for 90 min at room temperature and developed using western blotting reagents (BL523B, Bioshark, China). GAPDH (AF2823, Byotime, China) was used as the internal control.

Software and statistical analysis

All experiments were repeated at least once, and sample sizes and reported results reflect the cumulative data for all trials of each experiment. Each result is expressed as the mean ± standard deviation (SD). Pearson’s test was used to assess the normality of the data. The normally distributed data were subjected to parametric analyses. Unpaired Student’s t-tests were used for parametric analyses between two groups. Data that were not normally distributed or had a sample size that was too small, were subjected to nonparametric analyses. Nonparametric tests between two groups were performed using the Wilcoxon signed-rank test for matched pairs, the Mann-Whitney U-test or the Kruskal-Wallis test for multiple comparisons (GraphPad Prism 9, La Jolla, CA). A difference was considered statistically significant when the p-value was less than 0.05 and highly significant if it was less than 0.01 or less than 0.001.